Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

# Lab 2: Pandas Overview

[Pandas](https://pandas.pydata.org/) is one of the most widely used Python libraries in data science. In this lab, you will learn commonly used data wrangling operations and tools in Pandas. We aim to give you familiarity with:

* Creating dataframes
* Slicing data frames (i.e., selecting rows and columns)
* Filtering data (using boolean arrays)
* Data Aggregation and Grouping in dataframes

In this lab, you are going to use several pandas methods like `drop()`, `loc()`, `groupby()`. Remember that you can press `shift+tab` on any method to see the documentation for that method.

## Setup

```python
# Import the following packages. Note the shorthand for pandas.
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
%matplotlib inline
```

## Part 1: Creating DataFrames & Basic Manipulations

A [dataframe](https://pandas.pydata.org/pandas-docs/stable/user_guide/dsintro.html#dataframe) is a two-dimensional labeled data structure with columns holding data of potentially different types.

**Method 1: ** You can create a dataframe by specifying the columns and values as shown below.

Notice the syntax: you're passing a dictionary into the DataFrame.  The keys become the column names (e.g. `'name'`), and the values are lists (`['Peter',....`)

```python
animals = pd.DataFrame(
    data={'name': ['Peter', 'Nutkin', 'Hunca Munca', 'Jemima'],
          'species': ['rabbit', 'squirrel', 'mouse', 'duck']
          })
animals
```

Output:
```
name   species
0        Peter    rabbit
1       Nutkin  squirrel
2  Hunca Munca     mouse
3       Jemima      duck
```

**Method 2: ** You can also define a dataframe by specifying the rows like below.

Here, you're passing in tuples for each row of data (e.g. `("Peter", "rabbit")`) and specifying the column names separately.

```python
animals2 = pd.DataFrame(
    [("Peter", "rabbit"), ("Nutkin", "squirrel"), ("Hunca Munca", "mouse"),
     ("Jemima", "duck")], 
    columns = ["name", "species"])
animals2
```

Output:
```
name   species
0        Peter    rabbit
1       Nutkin  squirrel
2  Hunca Munca     mouse
3       Jemima      duck
```

**Other methods**: Usually you won't be creating data frames in such a manual way.  You'll often be loading dataframes in from other file types -- for example, comma separated (csv) files.  More on that later.

You can obtain the dimensions of a dataframe by using the shape attribute, `dataframe.shape`

```python
(num_rows, num_columns) = animals.shape
num_rows, num_columns
```

Output:
```
(4, 2)
```

### Question 1

You can add a column using the syntax `dataframe['new column name'] = [data]`. Add a column called `favorite food` to the `animals` table which contains the strings 'nut', 'carrot', 'corn', and 'cheese'. Use your best guess as to which animal prefers which food. 

<!-- grade_id: q1_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Solution
animals['favorite food'] = ['carrot','nut','cheese','corn']
```

```python
animals
```

Output:
```
name   species favorite food
0        Peter    rabbit        carrot
1       Nutkin  squirrel           nut
2  Hunca Munca     mouse        cheese
3       Jemima      duck          corn
```

### Question 2

Use the `.drop()` method to [drop](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.drop.html) the `favorite food` column you created, and save the new dataframe (without the `favorite food` column) to `animals_original`. Some notes:

* You'll need to look up `drop` to figure out the right syntax.
* Make sure to use the `axis` parameter correctly

<!-- grade_id: q2_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Solution 
animals_original = animals.drop(['favorite food'], axis=1)
```

```python
animals_original
```

Output:
```
name   species
0        Peter    rabbit
1       Nutkin  squirrel
2  Hunca Munca     mouse
3       Jemima      duck
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
assert animals_original.shape[1] == 2
```

### Question 3

Use the `.rename()` method to [rename](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.rename.html) the columns of `animals_original` so they begin with a capital letter. Set the `inplace` parameter correctly to change the `animals_original` dataframe. (hint: in Question 2, `drop` creates and returns a new dataframe instead of changing `animals` because `inplace` by default is `False`)

<!-- grade_id: q3_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Solution
animals_original.rename(columns = {"name": "Name", "species":"Species"}, inplace = True)
```

```python
animals_original
```

Output:
```
Name   Species
0        Peter    rabbit
1       Nutkin  squirrel
2  Hunca Munca     mouse
3       Jemima      duck
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
assert animals_original.columns[1] == 'Species' # the column number might be different for you
```

*Background*: For the curious, the field values you just worked with were inspired by [Beatrix Potter's](https://en.wikipedia.org/wiki/Beatrix_Potter) characters.

## Part 2: CalEnviroScreen Data
Now that we have learned the basics, we'll use Pandas to wrangle a real-world dataset. Specifically, we will be working with the [California Communities Environmental Health Screening Tool (CalEnviroScreen)](https://oehha.ca.gov/calenviroscreen), which uses demographic and environmental information to identify communities that are susceptible to various types of pollution. The various fields in this dataset contribute to the CES score, which reflects a community's environmental conditions and its vulnerability to environmental pollutants.

Your lab02 folder contains an Excel file downloaded from [here](https://oehha.ca.gov/calenviroscreen/report/draft-calenviroscreen-40).

Start by running the cell below, which creates an Excel file object in Pandas that we can then inspect. The cell below shows you the sheet names in the spreadsheet. Documentation on Pandas' Excel methods can be found [here](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.read_excel.html). 

```python
# run this cell to import the Excel file and see the names of the tabs
filename = 'CalEnviroScreen_4.0Excel_ADA_D1_2021.xlsx'
xl = pd.ExcelFile(filename)
print(xl.sheet_names) # display a list of the sheets in the spreadsheet
```

Output:
```
['DRAFTCES4.0_results', 'Demographic profile', 'Data Dictionary']
```

Run the cell below to load the first sheet of the Excel file and assign it to the variable `ces4`. 

```python
ces4 = xl.parse(xl.sheet_names[0]) # display the first sheet as Pandas dataframe
ces4.head()
```

Output:
```
Census Tract   Total Population California County    ZIP  \
0    6019001100               2760           Fresno   93706   
1    6077000700               4177       San Joaquin  95206   
2    6077000100               4055       San Joaquin  95202   
3    6071001600               5527    San Bernardino  91761   
4    6037204920               2639       Los Angeles  90023   

  Nearby City \n(to help approximate location only)   Longitude   Latitude  \
0                                            Fresno -119.781696  36.709695   
1                                          Stockton -121.287873  37.943173   
2                                          Stockton -121.285363  37.953996   
3                                           Ontario -117.618013  34.057780   
4                                       Los Angeles -118.197497  34.017500   

   DRAFT CES 4.0 Score  DRAFT CES 4.0 Percentile  \
0            94.608491                100.000000   
1            90.831885                 99.987390   
```

Note that the dataframe contains 58 columns, but Pandas truncates the number we are able to see at once. We can show all columns using the `pd.set_option` method.

```python
pd.set_option('display.max_columns', None)
ces4.head()
```

Output:
```
Census Tract   Total Population California County    ZIP  \
0    6019001100               2760           Fresno   93706   
1    6077000700               4177       San Joaquin  95206   
2    6077000100               4055       San Joaquin  95202   
3    6071001600               5527    San Bernardino  91761   
4    6037204920               2639       Los Angeles  90023   

  Nearby City \n(to help approximate location only)   Longitude   Latitude  \
0                                            Fresno -119.781696  36.709695   
1                                          Stockton -121.287873  37.943173   
2                                          Stockton -121.285363  37.953996   
3                                           Ontario -117.618013  34.057780   
4                                       Los Angeles -118.197497  34.017500   

   DRAFT CES 4.0 Score  DRAFT CES 4.0 Percentile  \
0            94.608491                100.000000   
1            90.831885                 99.987390   
```

Notice that this dataset doesn't include the units in many of the column headings. Let's load a different sheet to get more information about what we're looking at.

Run the following cell to load the data dictionary.

```python
dd = xl.parse('Data Dictionary', header = 6)
dd.head(10)
```

Output:
```
Longitude  \
0                        Latitude   
1             Draft CES 4.0 Score   
2        Draft CES 4.0 Percentile   
3  Draft CES 4.0 Percentile Range   
4                           Ozone   
5                      Ozone Pctl   
6                           PM2.5   
7                      PM2.5 Pctl   
8                       Diesel PM   
9                  Diesel PM Pctl   

       Longitude of the centroid of the census tract  
0       Latitude of the centroid of the census tract  
1  CalEnviroScreen Score, Pollution Score multipl...  
2            Percentile of the CalEnviroScreen score  
3  Percentile of the CalEnviroScreen score, group...  
4  Amount of daily maximum 8-hour Ozone concentra...  
5                                   Ozone percentile  
6                  Annual mean PM 2.5 concentrations  
7                                   PM2.5 percentile  
8  Diesel PM emissions from on-road and non-road ...  
9                               Diesel PM percentile
```

### Question 4
The length of a dataframe is equivalent to its number of rows. Find the length of `ces4`. What does each row represent?

```python
# YOUR CODE HERE
```

*YOUR ANSWER HERE*

```python
# Solution
len(ces4)
```

Output:
```
8035
```

*SOLUTION:* *Each row represents observations for a unique census tract.*

## Slicing Data Frames - selecting rows and columns


### Selection Using Label

**Column Selection** 
To select a column of a `DataFrame` by column label, the safest and fastest way is to use the `.loc` [method](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.loc.html). General usage looks like `frame.loc[rowname,colname]`. (Reminder that the colon `:` means "everything").

- You can also slice across columns. For example, `ces3.loc[:, 'ZIP':]` would select all rows in the column `ZIP` and every column to the right.

- *Alternative:* While `.loc` is invaluable when writing production code, it may be a little too verbose for interactive use. One recommended alternative is the `[]` method, which takes on the form `frame['colname']`.

**Row Selection**
Similarly, if we want to select a row by its label, we can use the same `.loc` method. In this case, the "label" of each row refers to the index (i.e., primary key) of the dataframe.

```python
#Example:
ces4.loc[100:110, 'ZIP']
```

Output:
```
100    90059
101    90201
102    90280
103    90280
104    92501
105    95203
106    93307
107    90023
108    95206
109    90222
110    90023
Name: ZIP, dtype: int64
```

```python
#Example:  Notice the difference between these two methods
ces4.loc[100:110, ['ZIP']]
```

Output:
```
ZIP
100  90059
101  90201
102  90280
103  90280
104  92501
105  95203
106  93307
107  90023
108  95206
109  90222
110  90023
```

The `.loc` method actually uses the index (the bolded, leftmost series in the dataframe) rather than the row position to perform the selection. In the previous example, it's just a coincidence that the `.loc` syntax matches that of the array slicing syntax - the index and row position aren't always the same value. For example, you could set your index to a non-numeric code, like census tract or other unique ID, if that's how you want to identify your records.

Alternatively, we can use [`.iloc`](http://pandas.pydata.org/pandas-docs/version/0.17.0/generated/pandas.DataFrame.iloc.html) to slice the dataframe using row location and column position.

See the following example:

```python
#Example: We change the index from 0,1,2... to the Census Tract column
df = ces4.set_index("Census Tract") # Why might we want to use Census Tract instead of County or City?
df.head()
```

Output:
```
Total Population California County    ZIP  \
Census Tract                                               
6019001100                 2760           Fresno   93706   
6077000700                 4177       San Joaquin  95206   
6077000100                 4055       San Joaquin  95202   
6071001600                 5527    San Bernardino  91761   
6037204920                 2639       Los Angeles  90023   

             Nearby City \n(to help approximate location only)   Longitude  \
Census Tract                                                                 
6019001100                                              Fresno -119.781696   
6077000700                                            Stockton -121.287873   
6077000100                                            Stockton -121.285363   
6071001600                                             Ontario -117.618013   
6037204920                                         Los Angeles -118.197497   

               Latitude  DRAFT CES 4.0 Score  D
```

We can now lookup rows by name directly:

```python
df.loc[[6037205120, 6019000200], :]
```

Output:
```
Total Population California County    ZIP  \
Census Tract                                               
6037205120                 3521       Los Angeles  90023   
6019000200                 2870           Fresno   93706   

             Nearby City \n(to help approximate location only)   Longitude  \
Census Tract                                                                 
6037205120                                         Los Angeles -118.211796   
6019000200                                              Fresno -119.805504   

               Latitude  DRAFT CES 4.0 Score  DRAFT CES 4.0 Percentile  \
Census Tract                                                             
6037205120    34.018755            74.556121                 99.659521   
6019000200    36.735491            79.894464                 99.899117   

             DRAFT CES 4.0\nPercentile Range     Ozone  Ozone Pctl      PM2.5  \
Census Tract                                                                    
60
```

However, if we want to access rows by location we will need to use the integer loc (`iloc`) accessor:

```python
#Example: 
# df.loc[2:5,"Year"] # You can't do this
df.iloc[1:4,6:9]
```

Output:
```
DRAFT CES 4.0 Score  DRAFT CES 4.0 Percentile  \
Census Tract                                                  
6077000700              90.831885                 99.987390   
6077000100              85.750324                 99.974779   
6071001600              83.556404                 99.962169   

             DRAFT CES 4.0\nPercentile Range  
Census Tract                                  
6077000700          95-100% (highest scores)  
6077000100          95-100% (highest scores)  
6071001600          95-100% (highest scores)
```

### Question 5

Selecting multiple columns is easy using `.loc`.  You just need to supply a list of column names.  Select the `California County`,`Diesel PM Pctl`, and `PM2.5 Pctl` columns **in that order** from the `ces3` table.

<!-- grade_id: q4_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Solution
dsl_PM = ces4.loc[:, ["California County", "Diesel PM Pctl","PM2.5 Pctl"]]
```

```python
dsl_PM.head()
```

Output:
```
California County  Diesel PM Pctl  PM2.5 Pctl
0           Fresno        98.718108   97.710019
1       San Joaquin       91.151213   72.594897
2       San Joaquin       90.628500   72.196640
3    San Bernardino       86.011201   97.398880
4       Los Angeles       96.552582   89.209708
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
assert dsl_PM.shape == (8035, 3)
assert dsl_PM.columns[1] == "Diesel PM Pctl"
```

As you may have noticed above, the .loc() method is a way to re-order the columns within a dataframe.

## Filtering Data

### Filtering with boolean arrays

Filtering is the process of removing unwanted material.  In your quest for cleaner data, you will undoubtedly filter your data at some point: whether it be for clearing up cases with missing values, culling out fishy outliers, or analyzing subgroups of your data set.  Note that compound expressions have to be grouped with brackets. Example usage looks like `df[df[column name] < 5]]`.

For your reference, some commonly used comparison operators are given below.

Symbol | Usage      | Meaning 
------ | ---------- | -------------------------------------
==   | a == b   | Does a equal b?
<=   | a <= b   | Is a less than or equal to b?
>=   | a >= b   | Is a greater than or equal to b?
<    | a < b    | Is a less than b?
&#62;    | a &#62; b    | Is a greater than b?
~    | ~p       | Returns negation of p
&#124; | p &#124; q | p OR q
&    | p & q    | p AND q
^  | p ^ q | p XOR q (exclusive or)

In the following we construct the DataFrame containing only census tracts in Sacramento County.

```python
ces4_SC = ces4[ces4['California County'] == 'Sacramento '] # Note the space after "Sacramento." This kind of quirk can be remedied with some simple data cleaning techniques, which we'll discuss in future lessons.
ces4_SC.head()
```

Output:
```
Census Tract   Total Population California County    ZIP  \
126    6067005301               1462       Sacramento   95811   
379    6067006400               5686       Sacramento   95838   
424    6067000700               2604       Sacramento   95814   
581    6067004502               5472       Sacramento   95823   
618    6067006900               4626       Sacramento   95815   

    Nearby City \n(to help approximate location only)   Longitude   Latitude  \
126                                        Sacramento -121.482520  38.591604   
379                                        Sacramento -121.420238  38.641645   
424                                        Sacramento -121.501660  38.581871   
581                                        Sacramento -121.462926  38.502552   
618                                        Sacramento -121.457865  38.601870   

     DRAFT CES 4.0 Score  DRAFT CES 4.0 Percentile  \
126            66.118375                 98.411097   
379            58.603288 
```

### Question 6
Select the census tracts in Alameda county whose CES 3.0 Percentile is 90 or higher.

(If you use condition `p` & condition `q` to filter the dataframe, make sure to use `df[(p) & (q)]`)

Hint: The county names are not "clean." Try using the `.unique()` method to look up the **exact** county names. 

<!-- grade_id: q5a_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Solution
AC_highCES = ces4[(ces4["California County"] == "Alameda ") & (ces4['DRAFT CES 4.0 Percentile']>=90)]
```

```python
AC_highCES
```

Output:
```
Census Tract   Total Population California County    ZIP  \
162    6001407300               2440          Alameda   94601   
200    6001408800               7054          Alameda   94621   
217    6001409000               3904          Alameda   94621   
331    6001409100               2327          Alameda   94603   
575    6001403000               2969          Alameda   94607   
577    6001402200               2527          Alameda   94607   
587    6001406100               4659          Alameda   94601   
715    6001409500               4272          Alameda   94621   
762    6001406000               3344          Alameda   94606   
767    6001410500               2373          Alameda   94607   

    Nearby City \n(to help approximate location only)   Longitude   Latitude  \
162                                           Oakland -122.210924  37.762179   
200                                           Oakland -122.196942  37.758804   
217                                           Oak
```

<!-- grade_id: q5a_test, points: 1, grade: true, solution: false -->
```python
assert len(AC_highCES) == 10
assert AC_highCES["DRAFT CES 4.0 Percentile"].max() == 97.9571248423707
```

## Data Aggregration 

### Question 7a
We can perform operations across columns and rows of a dataframe to generate summary statistics. 

Find the mean, minimum, maximum, and standard deviation PM2.5 across all census tracts. (You may use the `ces4` DataFrame created above.)

```python
# YOUR CODE HERE

pm_mean = ...
pm_min = ...
pm_max = ...
pm_stdev = ...

print('Mean: {}'.format(pm_mean))
print('Min: {}'.format(pm_min))
print('Max: {}'.format(pm_max))
print('Standard deviation: {}'.format(pm_stdev))
```

Output:
```
Mean: Ellipsis
Min: Ellipsis
Max: Ellipsis
Standard deviation: Ellipsis
```

```python
# Solution
pm_mean = ces4['PM2.5'].mean()
pm_min = ces4['PM2.5'].min()
pm_max = ces4['PM2.5'].max()
pm_stdev = ces4['PM2.5'].std()

print('Mean: {}'.format(pm_mean))
print('Min: {}'.format(pm_min))
print('Max: {}'.format(pm_max))
print('Standard deviation: {}'.format(pm_stdev))
```

Output:
```
Mean: 10.152699916179465
Min: 1.875091539
Max: 16.39474779
Standard deviation: 2.165937022827508
```

### Question 7b

What is the total population reprepresented in the CalEnviroscreen 4.0 dataset?

```python
# YOUR CODE HERE
```

```python
# Solution
ces4[' Total Population'].sum()
```

Output:
```
39148760
```

### Question 8
To count the number of instances of a value in a `Series`, we can use the `value_counts()` [method](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.Series.value_counts.html) as `df["col_name"].value_counts()`. 

Count the number of different census tracts in each California county. (You may use the `ces4` DataFrame created above.) In other words, compute the number of rows in the table for each county.

<!-- grade_id: q6a_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Solution

num_censusTracts = ces4["California County"].value_counts()
```

```python
num_censusTracts
```

Output:
```
Los Angeles        2343
San Diego           627
Orange              582
Riverside           453
Santa Clara         372
San Bernardino      369
Alameda             360
Sacramento          317
Contra Costa        207
Fresno              199
San Francisco       195
Ventura             173
San Mateo           157
Kern                151
San Joaquin         139
Sonoma               99
Solano               96
Stanislaus           94
Monterey             93
Santa Barbara        89
Placer               84
Tulare               78
Marin                55
San Luis Obispo      53
Santa Cruz           52
Butte                51
Merced               49
Shasta               48
El Dorado            42
Yolo                 41
Napa                 40
Imperial             31
Humboldt             30
Kings                27
Madera               23
Sutter               21
Mendocino            20
Nevada               20
Lake                 15
Yuba                 14
Siskiyou             14
San Benito      
```

<!-- grade_id: q6a_test, points: 1, grade: true, solution: false -->
```python
assert num_censusTracts["Alameda "] == 360
assert num_censusTracts.sum() == len(ces4)
```

### Question 9a

A more versatile way to aggregate data is to use the `.groupby()` [function](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.groupby.html). Find the sum of `Imp. Water Bodies` for each `California County` in the `ces4` table. Use the syntax `df.groupby("col_name").sum()`.

<!-- grade_id: q7a_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Solution
sum_impH2O = ces4.groupby("California County").sum()
```

```python
# The syntax below will sort sum_impH20 from most to least impaired bodies. 
sum_impH2O.sort_values(by = "Imp. Water Bodies", ascending=False)['Imp. Water Bodies']
```

Output:
```
California County
Los Angeles        7178
San Diego          3617
Orange             1742
Ventura            1493
Alameda            1461
Sacramento         1344
San Joaquin        1074
Contra Costa       1044
San Francisco      1043
Santa Clara         838
Monterey            809
Riverside           668
Santa Barbara       612
San Mateo           596
Santa Cruz          504
Sonoma              502
Imperial            497
Stanislaus          473
Solano              362
Marin               346
San Luis Obispo     337
San Bernardino      284
Placer              261
Merced              225
Yolo                210
Sutter              195
Butte               170
Fresno              129
Humboldt            125
Tulare              104
El Dorado           100
Shasta               98
Napa                 86
San Benito           86
Siskiyou             77
Nevada               70
Mendocino            69
Colusa               66
Madera               59
Yuba                 58
Lake                 5
```

```python
assert sum_impH2O.loc["Los Angeles", "Imp. Water Bodies"] == 7178
assert sum_impH2O.sort_values(by = "Imp. Water Bodies", ascending=False).index[3] == "Ventura "
```

### Question 9b
Take a look at the the Data Dictionary. What does the sum of the `Imp. Water Bodies` column represent?  In the process you'll read about "buffers".  What is a buffer?

```python
# SCRATCH WORK HERE
```

*YOUR ANSWER HERE*

*SOLUTION*

*The sum of the number of pollutants across all impaired water bodies within buffered distances to populated blocks in the **county**.* 

*Buffers are areas around an object in space (often a point).  They are often used to aggregate data (as in `groupby`: by averaging, taking the max, counts, etc.) in the vicinity of the object.*

### Question 9c

What do the values in `ZIP` represent in the dataframe `sum_impH2O`? Why is the column `DRAFT CES 4.0 Percentile Range` no longer present in the dataframe?

*YOUR ANSWER HERE*

*SOLUTION*

*`ZIP` is the sum of all zip codes associated with the census tracts in a given county. Although this isn't particularly meaningful to us, it's how Python interprets the `.sum()` method of `groupby()` - it takes the sum of all numeric values in the dataframe, grouped by the specified column (`California County`).*

*`CES 3.0 Percentile Range` is no longer present in the dataframe because the column does not contain numeric values. Python isn't able to sum the column's values, so it drops it. We could keep this column by grouping by both `California County` and `CES 3.0 Percentile Range` in the groupby - that way, we would get a dataframe that has the total number of pollutants in impaired water bodies by both county and CES 3.0 category.*

### Question 9d

Find the mean of `Poverty` for each county for census tracts with population greater than or equal to 3,000 and with a `Pollution Burden Pctl` above 85.


<!-- grade_id: q7b_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# SOLUTION
poverty_mean = ces4[(ces4[" Total Population"] >= 3000) & (ces4["Pollution Burden Pctl"] > 85)].groupby("California County").mean()
```

```python
poverty_mean.sort_values(by = "Poverty", ascending=False)['Poverty']
```

Output:
```
California County
Sutter            70.700000
Yuba              59.400000
Butte             59.400000
Fresno            59.302778
Tulare            57.381818
Kern              55.966667
Kings             53.175000
Imperial          52.400000
Yolo              50.000000
Madera            49.366667
Santa Barbara     46.250000
San Joaquin       45.865385
Merced            45.655556
Riverside         45.458974
San Diego         45.310000
Los Angeles       43.803061
Stanislaus        43.366667
Sacramento        42.540000
San Bernardino    42.179710
Monterey          40.900000
Santa Cruz        36.000000
Orange            35.909890
Placer            33.700000
Alameda           33.385714
Solano            32.566667
San Francisco     31.316667
Ventura           29.483333
Santa Clara       28.766667
Contra Costa      28.200000
San Mateo         22.966667
Name: Poverty, dtype: float64
```

<!-- grade_id: q7b_test, points: 1, grade: true, solution: false -->
```python
assert np.round(poverty_mean.loc["Alameda ", "Poverty"],2) == 33.39
assert len(poverty_mean) == 30
```

### Question 9e

What does your output to 7d represent?  Dig in to the data a little further and tell us what you notice about the `Poverty` field values in the counties/tract combinations that show up in your result to 7d, versus the `Poverty` field values for all tracts?

*YOUR ANSWER HERE*

*From the CES documentation, poverty is the "Percent of the population living below two times the federal poverty
level (5-year estimate, 2011-2015)."  If one digs in to the data, you can find that the poverty percentile for those tracts with high CES scores is higher than the average percentile across all tracts.*  

#### You are done! Remember to submit this lab on bCourses in both html and ipynb formats after clicking Kernel -> Restart & Run All.